---------------------------------------------------------------------------
SZS status Theorem for test/problems/prop-040.tptp

SZS output start CNFRefutation for test/problems/prop-040.tptp
fof(goal, conjecture,
    ((~ a & (a | b) & (c | d) & (~ b | e | f) & (~ c | ~ e) & (~ c | ~ f) &
      (~ b | g | h) & (~ d | ~ g) & (~ d | ~ h)) => $false)).

fof(subgoal_0, plain,
    ((~ a & (a | b) & (c | d) & (~ b | e | f) & (~ c | ~ e) & (~ c | ~ f) &
      (~ b | g | h) & (~ d | ~ g) & (~ d | ~ h)) => $false),
    inference(strip, [], [goal])).

fof(negate_0_0, plain,
    (~
       ((~ a & (a | b) & (c | d) & (~ b | e | f) & (~ c | ~ e) &
         (~ c | ~ f) & (~ b | g | h) & (~ d | ~ g) & (~ d | ~ h)) =>
        $false)), inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (~ a & (~ c | ~ e) & (~ c | ~ f) & (~ d | ~ g) & (~ d | ~ h) &
     (a | b) & (c | d) & (~ b | e | f) & (~ b | g | h)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain, (~ c | ~ e),
    inference(conjunct, [], [normalize_0_0])).

fof(normalize_0_2, plain, (~ c | ~ f),
    inference(conjunct, [], [normalize_0_0])).

fof(normalize_0_3, plain, (~ b | e | f),
    inference(conjunct, [], [normalize_0_0])).

fof(normalize_0_4, plain, (a | b),
    inference(conjunct, [], [normalize_0_0])).

fof(normalize_0_5, plain, (~ a), inference(conjunct, [], [normalize_0_0])).

fof(normalize_0_6, plain, (b),
    inference(simplify, [], [normalize_0_4, normalize_0_5])).

fof(normalize_0_7, plain, (e | f),
    inference(simplify, [], [normalize_0_3, normalize_0_6])).

fof(normalize_0_8, plain, (c | d),
    inference(conjunct, [], [normalize_0_0])).

fof(normalize_0_9, plain, (~ d | ~ g),
    inference(conjunct, [], [normalize_0_0])).

fof(normalize_0_10, plain, (~ d | ~ h),
    inference(conjunct, [], [normalize_0_0])).

fof(normalize_0_11, plain, (~ b | g | h),
    inference(conjunct, [], [normalize_0_0])).

fof(normalize_0_12, plain, (g | h),
    inference(simplify, [], [normalize_0_11, normalize_0_6])).

cnf(refute_0_0, plain, (~ c | ~ e),
    inference(canonicalize, [], [normalize_0_1])).

cnf(refute_0_1, plain, (~ c | ~ f),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_2, plain, (e | f),
    inference(canonicalize, [], [normalize_0_7])).

cnf(refute_0_3, plain, (~ c | e),
    inference(resolve, [$cnf(f)], [refute_0_2, refute_0_1])).

cnf(refute_0_4, plain, (c | d),
    inference(canonicalize, [], [normalize_0_8])).

cnf(refute_0_5, plain, (~ d | ~ g),
    inference(canonicalize, [], [normalize_0_9])).

cnf(refute_0_6, plain, (~ d | ~ h),
    inference(canonicalize, [], [normalize_0_10])).

cnf(refute_0_7, plain, (g | h),
    inference(canonicalize, [], [normalize_0_12])).

cnf(refute_0_8, plain, (~ d | g),
    inference(resolve, [$cnf(h)], [refute_0_7, refute_0_6])).

cnf(refute_0_9, plain, (c | g),
    inference(resolve, [$cnf(d)], [refute_0_4, refute_0_8])).

cnf(refute_0_10, plain, (~ d | c),
    inference(resolve, [$cnf(g)], [refute_0_9, refute_0_5])).

cnf(refute_0_11, plain, (c),
    inference(resolve, [$cnf(d)], [refute_0_4, refute_0_10])).

cnf(refute_0_12, plain, (e),
    inference(resolve, [$cnf(c)], [refute_0_11, refute_0_3])).

cnf(refute_0_13, plain, (~ c),
    inference(resolve, [$cnf(e)], [refute_0_12, refute_0_0])).

cnf(refute_0_14, plain, ($false),
    inference(resolve, [$cnf(c)], [refute_0_11, refute_0_13])).
SZS output end CNFRefutation for test/problems/prop-040.tptp

